// variables5.rs
//
// Execute `rustlings hint variables5` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE
// TODO
fn main() {
    let number: &str = "T-H-R-E-E"; // don't change this line
    // change the number type to str
    
    println!("Spell a Number : {}", number);
    // number = 3; // don't rename this variable
    println!("Number plus two is : {}", number.to_string().len() + 2);
}
