// primitive_types1.rs
//
// Fill in the rest of the line that has code missing! No hints, there's no
// tricks, just get used to typing these :)

// I AM NOT DONE

fn main() {
    // Booleans (`bool`)

    let is_morning: bool = true;
    if is_morning {
        println!("Good morning!");
    }

    let is_evening: bool = false; // Declare and assign a boolean value to is_evening
    if is_evening {
        println!("Good evening!");
    }
}
